# <div style="background-color:rgba(204, 229, 255, 0.5); text-align:center; vertical-align: middle; padding:40px 0; margin-top:30px"><span style="color:rgba(0, 76, 153, 1);">PHYS 231 Lab #4 - Day 2<span style="color:red"> $\to$ ($13$ possible marks)</span></span></div>

### Due Thursday, Oct. 26 at 17:00 $-$ upload your .zip submission to the PHYS 231 Canvas Gradebook.

<font size ="4"><span style="color:purple">We recommend that you complete this notebook using either the [Google Chrome](https://www.google.com/intl/en_ca/chrome/) or [Mozilla Firefox](https://www.mozilla.org/en-CA/firefox/new/) browser.  Chrome and Firefox are both available for Windows, Mac, and Linux operating systems.  The PHYS 121 Jupyter notebooks should work in other browsers, but we have specifically verfied that they work well in both Chrome and Firefox.</span></font>

# $LRC$ Circuit Frequency Response

***
## Autograding:
The PHYS 231 Labs will make use of some autograding.  To make the autograding work, the two cells below needs to be executed.  The first one installs the required packages and the second imports required packages/modules.  If 'PHYS231.Installer()' reports that some functions have been installed, the user should restart the "kernel" and then exectute the PHYS231.Installer() cell a second time.  The second time the installer function is run, it should report that **"All packages already installed. Please proceed"**.

If necessary, the kernel can be restarted by selecting **Kernel** $\to$ **Restart Kernel** from the menu options at to the top of the page.  Here is a <a href = "https://raw.githubusercontent.com/UBC-Okanagan-Physics-Labs/PHYS-121-images/main/general/gifs/restartKernel.gif">GIF</a> showing how to restart the kernel.

The 'PHYS231.Installer()' command requires the file 'PHYS231.py', which you should see included in the list of files along the left-hand side of the screen.

```python
# Import PHYS121.py and then run the installer function.
import PHYS231
PHYS231.Installer()
```

```python
# Initialize Otter
import otter
grader = otter.Notebook("PHYS 231 - Lab 4b.ipynb")
```

***
## Import Modules:
Execute the cell below to import a number of useful pre-built Python modules that will be used in the PHYS 121 pre-labs and labs.

```python
import numpy as np
import matplotlib.pyplot as plt
from matplotlib.pyplot import cm # used to generate a sequence of colours for plotting
import pandas as pd
import scipy.stats as sp
from scipy.optimize import curve_fit
from IPython.display import HTML as html_print
from IPython.display import display, Markdown, Latex
import uncertainties
```

# <div style="background-color:rgba(255, 204, 255, 0.5); text-align:center; vertical-align: middle; padding:40px 0; margin-top:30px"><span style="color:rgba(102, 0, 204, 1);">Name & Student Number</span></div>  

To ensure that both you and your partner receive credit for you lab work, both your name and student number and your partner's name and student number need to appear in the notebooks used during the labs.  In the cells below this text, enter your name and student number and your partner's name and student number as strings (i.e. place your names and student numbers between single or double quotes.

Examples of valid entries are:
```python
last_name_1 = 'Doe'
first_name_1 = 'Jane'
student_number_1 = '12345678'
```
and
```python
last_name_1 = "Doe"
first_name_1 = "John"
student_number_1 = "12345678"
```

<!-- BEGIN QUESTION -->

***
### **<span style="color:blue">Partner #1</span>** **<span style="color:red">(0.5 marks)</span>** 

```python
# Enter LAST name of partner 1.  Delete the ... below and type your entry between the single quotes.
# Once you've completed your entry, press SHIFT + ENTER to execute the cell.
# *** Please do not change anything to the left of the equals sign. ***
last_name_1 = ...
```

```python
# Enter FIRST name of partner 1.  Delete the ... below and type your entry between the single quotes.
# Once you've completed your entry, press SHIFT + ENTER to execute the cell.
# *** Please do not change anything to the left of the equals sign. ***
first_name_1 = ...
```

```python
# Enter STUDENT NUMBER of partner 1.  Delete the ... below and type your entry between the single quotes.
# Once you've completed your entry, press SHIFT + ENTER to execute the cell.
# *** Please do not change anything to the left of the equals sign. ***
student_number_1 = ...
```

***
### **<span style="color:blue">Partner #2</span>** **<span style="color:red">(0.5 marks)</span>** 
If you are working on your own, leave all of last_name_2, first_name_2, and student_number_2 as ...  Please still exectute these cells and then run the 'grader.check("name_and_student_number_2")' cell.

```python
# Enter LAST name of partner 2.  Delete the ... below and type your entry between the single quotes.
# Once you've completed your entry, press SHIFT + ENTER to execute the cell.
# *** Please do not change anything to the left of the equals sign. ***
last_name_2 = ...
```

```python
# Enter FIRST name of partner 2.  Delete the ... below and type your entry between the single quotes.
# Once you've completed your entry, press SHIFT + ENTER to execute the cell.
# *** Please do not change anything to the left of the equals sign. ***
first_name_2 = ...
```

```python
# Enter STUDENT NUMBER of partner 2.  Delete the ... below and type your entry between the single quotes.
# Once you've completed your entry, press SHIFT + ENTER to execute the cell.
# *** Please do not change anything to the left of the equals sign. ***
student_number_2 = ...
```

```python
grader.check("name_and_student_number")
```

<!-- END QUESTION -->

# <div style="background-color:rgba(255, 204, 255, 0.5); text-align:center; vertical-align: middle; padding:40px 0; margin-top:30px"><span style="color:rgba(102, 0, 204, 1);">Part 1 - Oscilloscope Measurement of the Resonance and Phase</span></div>

## NumPy Arrays

Data that you collect during the lab can be entered directly into a list using the following notation:

```python
x = [x0, x1, x2, x3, ..., xN]
y = [y0, y1, y2, y3, ..., yN]
```
These lists will contain $N+1$ elements. It's usually more convenient to enter data into a NumPy array rather than a list because you can do mathematical operations on arrays that you can't do on lists.  For example:

- squaring an array creates an new array with each of the elements from the original array squared
- multiplying two arrays of the same length creates a new array in which the each elements is the product of the corresponding elements from the original arrays
- ...

NumPy arrays are entered using the following syntax:

```python
x = np.array([x0, x1, x2, x3, ..., xN])
y = np.array([y0, y1, y2, y3, ..., yN])
```             

In the cell below, replace the "..." in NumPy array statements with your measurements of:

- frequency in kHz
- the amplitude of the signal from the function generator V0 in volts
- the uncertainty in V0 in volts
- the amplitude of the voltage across the resistor VR in volts
- the uncertainty in VR in volts
- the difference in time in $\rm \mu s$ between V0 and VR at, say, a zero crossing with positive slope (note that some of the time differences will be positive and some will be negative)
- the uncertainty in the time difference in $\rm \mu s$

<!-- BEGIN QUESTION -->

***
**<span style="color:blue">Question 1.1</span>** **<span style="color:red">(2 marks)</span>**  

In the cell below, replace the ... with your data.  You have **lots** of time in today's lab, so focus on collecting as much high-quality data as possible.  Collect data at a **minimum** of $N = 25$ different frequencies.  Please enter your frequency data in **units of kHz** and your dt and errdt data in **units of $\mu s$**.

*** Please do not change anything to the left of the equals sign. ***

```python
frequency = np.array([..])
V0 = np.array([...])
errV0 = np.array([...])
VR = np.array([...])
errVR = np.array([...])
dt = np.array([...])
errdt = np.array([...])

print("Array Lengths")
print("frequency: ", len(frequency))
print("V0: ", len(V0))
print("errV0: ", len(errV0))
print("VR: ", len(VR))
print("errVR: ", len(errVR))
print("dt: ", len(dt))
print("errdt: ", len(errdt))
```

```python
grader.check("q1.1")
```

<!-- END QUESTION -->

## Pandas DateFrames

The Pandas package can be used to display your array data in a nicely-formated table.  The Pandas package creates "DataFrames" which are 2 dimensional data structures, like a 2 dimensional array, or a table with rows and columns.  The cell below creates and then displays a DataFrame using the arrays that you created in the cell above.

```python
# Create the data structure.
data = {
    "frequnecy (kHz)": frequency,
    "V0 (V)": V0,
    "errV0 (V)": errV0,
    "VR (V)": VR,
    "errVR (V)": errVR,
    "dt ($\mu$s)": dt,
    "errdt ($\mu$s)": errdt
        }

# An option to format the numbers in the DataFrame.
pd.set_option('display.float_format', lambda x: '%0.3g' % x)

#Display the data.
pd.DataFrame(data)
```

## $V_\mathrm{R}/V_0$ versus $\omega = 2\pi f$

In the cells below, calculate:

- $\omega = 2\pi f$.  For $\pi$, use the syntax ```np.pi```.  
- $V_\mathrm{R}/V_0$
- the uncertainty in $V_\mathrm{R}/V_0$

We'll assume the relative uncertainty in the frequency is negligible compared the relative uncertainties of the other relevant parameters.

Recall propagation of errors.  If $y = f\left(x1, x2, ..., x_N\right)$ and $x_i\pm\Delta x_i$ for $i= 1...N$ are known, then:

\begin{align}
\Delta y = \sqrt{\left(\frac{\partial f}{\partial x_1}\Delta x_1\right)^2 + \left(\frac{\partial f}{\partial x_2}\Delta x_2\right)^2 + ... + \left(\frac{\partial f}{\partial x_N}\Delta x_N\right)^2}
\end{align}

***
**<span style="color:blue">Question 1.2</span>** **<span style="color:red">(3 marks)</span>**  

In the cell below, replace the ... with expressions that complete the required calculations.

*** Please do not change anything to the left of the equals sign. ***

```python
# Calucate the angular frequency in units of krad/s.
omega = ...
Ratio = ...

# Note that, in Python, NumPy can be used to access the square root function.  The syntax is np.sqrt(...).
errRatio = ...
```

```python
grader.check("q1.2")
```

### Plot the data

The $V_\mathrm{R}/V_0$ data can now be plotted as a function of angular frequency.

<!-- BEGIN QUESTION -->

***
**<span style="color:blue">Question 1.3</span>** **<span style="color:red">(1 mark)</span>**  

In the cell below, replace the ... with the appropriate data to generate a plot of $V_R/V_0$ versus $\omega$.  Be sure to include errror bars.

For additional information about the 'Scatter()' function, see **Appendix A** near the end of this notebook. 

*** Please do not change anything to the left of the equals sign. ***

```python
PHYS231.Scatter(xData = ..., 
                yData = ..., 
                yErrors = ...,
                xlabel = '$\omega$', 
                ylabel = '$V_R/V_0$', 
                xUnits = 'krad/s', 
                yUnits = '');
```

<!-- END QUESTION -->

## Nonlinear Fits

Recall from class that, for the series $LRC$ circuit:

\begin{equation}
\frac{V_\mathrm{R}}{V_0} = \frac{R}{\sqrt{R^2 + \left(\omega L - \dfrac{1}{\omega C}\right)^2}}
\end{equation}

If the substitutions $\gamma = R/L$ and $\omega_0 = 1/\sqrt{LC}$ are made, one can show that:

\begin{equation}
\frac{V_\mathrm{R}}{V_0} = \frac{1}{\sqrt{1 + \left(\dfrac{\omega}{\gamma}\right)^2\left[1 - \left(\dfrac{\omega_0}{\omega}\right)^2\right]^2}}
\end{equation}

The parameter $\omega_0$ represents the resonance frequency and $\gamma$ represents the width of the resonance.

In practice, the peak of your resonance will not reach a value of $V_\mathrm{R}/V_0 = 1$.  This is due to losses on the wire used to wind the inductors which have not been taken into account.  To capture the effect of these losses, we can add an overall scaling factor $A$ such that:

\begin{equation}
\frac{V_\mathrm{R}}{V_0} = \frac{A}{\sqrt{1 + \left(\dfrac{\omega}{\gamma}\right)^2\left[1 - \left(\dfrac{\omega_0}{\omega}\right)^2\right]^2}}
\end{equation}

You will fit this function to your data in order to extract the best-fit parameters $A\pm\Delta A$, $\omega_0\pm\Delta\omega_0$, and $\gamma\pm\Delta\gamma$.  

For nonlinear functions, it is often necessary to provide initial "guesses" for the parameters which are used by the fitting algorithm.  Use the plot generated above to determine reasonable initial guesses for the parameters $A$, $\omega_0$, and $\gamma$.

<!-- BEGIN QUESTION -->

***
**<span style="color:blue">Question 1.4</span>** **<span style="color:red">(1 mark)</span>**  

In the cell below, replace the ... with initial estimates of for the parameters $A$, $\omega_0$, and $\gamma$.

*** Please do not change anything to the left of the equals sign. ***

```python
A_init = ...
omega0_init = ... # krad/s
gamma_init = ... # krad/s

start1 = np.array([.9, 220, 50])
```

<!-- END QUESTION -->

We'll use the 'Lorentz' function contained in PHYS231.py to complete the nonlinear fit:
```python
A, Omega0, Gamma, errA, errOmega0, errGamma, fig = PHYS231.Lorentz(x, y, erry)
```
A weighted fit with a formatted plot can be called by using some of the other available arguments:
```python
A, Omega0, Gamma, errA, errOmega0, errGamma, fig = PHYS231.Lorentz(xData = x, yData = y, yErrors = erry, start = start, xlabel = 'time', ylabel = 'position', xUnits = 's', yUnits = 'm')
```

<!-- BEGIN QUESTION -->

***
**<span style="color:blue">Question 1.5</span>** **<span style="color:red">(1 mark)</span>**  

In the cell below, replace the ... with appropriate data to complete the weighted nonlinear fit to $V_R/V_0\pm\Delta\left(V_R/V_0\right)$ versus $\omega$.

*** Please do not change anything to the left of the equals sign. ***

```python
A, Omega0, Gamma, errA, errOmega0, errGamma, fig2 = PHYS231.Lorentz(
    xData = ..., 
    yData = ..., 
    yErrors = ..., 
    start = start1,
    xlabel = "$\omega$", 
    ylabel = '$V_R/V_0$', 
    xUnits = 'krad/s', 
    yUnits = '')
```

<!-- END QUESTION -->

In your lab notebook, you should compare this experimentally-measured values of $\omega_0$ and $\gamma$ to the values that you expected based on the values of the components used in your circuit.  Recall that $\omega_0 = 1/\sqrt{LC}$ and $\gamma=R/L$.  Keep in mind that your expected values of $\omega_0$ and $\gamma$ will also have uncertainties based on the uncertainties in the values of $R\pm\Delta R$, $L\pm\Delta L$, and $C\pm\Delta C$.

## $\phi$ versus $\omega$

In the cells below, using $\omega$ and the measured time difference $dt$ to calculate the phase difference between $V_\mathrm{R}$ and $V_0$ (and its uncertainty).  Recall that $\omega$ is in $\rm krad/s$ and $dt$ is in $\rm \mu s$.  You will need to convert to $\rm rad/s$ and $\rm s$ to get the correct values for $\phi$.

Please do **NOT** change the variable names.

***
**<span style="color:blue">Question 1.6</span>** **<span style="color:red">(2 marks)</span>**  

In the cell below, replace the ... an expression to calculate $\phi$ from the $dt$ and frequency measurements.

*** Please do not change anything to the left of the equals sign. ***

```python
Phi = ...
errPhi = ...
```

```python
grader.check("q1.6")
```

### Plot the data

The $\phi$ data can now be plotted as a function of angular frequency.

<!-- BEGIN QUESTION -->

***
**<span style="color:blue">Question 1.7</span>** **<span style="color:red">(1 mark)</span>**  

In the cell below, replace the ... with the appropriate data to generate a plot of $\Phi$ versus $\omega$.  Be sure to include errror bars.

For additional information about the 'Scatter()' function, see **Appendix A** near the end of this notebook. 

*** Please do not change anything to the left of the equals sign. ***

```python
PHYS231.Scatter(xData = ..., 
                yData = ..., 
                yErrors = ...,
                xlabel = '$\omega$', 
                ylabel = '$\Phi$', 
                xUnits = 'krad/s', 
                yUnits = 'rad');
```

<!-- END QUESTION -->

## Nonlinear Fits

Recall from class that, for the series $LRC$ circuit:

\begin{equation}
\phi = \arctan\left(\dfrac{1}{\omega RC} - \omega\frac{L}{R}\right)
\end{equation}

Expressed in terms of $\gamma = R/L$ and $\omega_0 = 1/\sqrt{LC}$, the phase becomes:

\begin{equation}
\phi = \arctan\left\{ \frac{\omega}{\gamma}\left[\left(\frac{\omega_0}{\omega}\right)^2 - 1\right] \right\}
\end{equation}

You will fit this function to your data in order to extract the best-fit parameters $\omega_0\pm\Delta\omega_0$ and $\gamma\pm\Delta\gamma$.  

We'll use the 'Phase' function contained in PHYS231.py to complete the nonlinear fit:
```python
Omega0, Gamma, errOmega0, errGamma, fig = PHYS231.Phase(x, y, erry)
```
A weighted fit with a formatted plot can be called by using some of the other available arguments:
```python
Omega0, Gamma, errOmega0, errGamma, fig = PHYS231.Phase(xData = x, yData = y, yErrors = erry, start = start, xlabel = 'time', ylabel = 'position', xUnits = 's', yUnits = 'm')
```

<!-- BEGIN QUESTION -->

***
**<span style="color:blue">Question 1.8</span>** **<span style="color:red">(1 mark)</span>**  

In the cell below, replace the ... with appropriate data to complete the weighted nonlinear fit to $\phi\pm\Delta\phi$ versus $\omega$.

*** Please do not change anything to the left of the equals sign. ***

```python
# We'll use the best-fit parameters from the first fit as the starting parameters for the fit to the Phi data.
start2 = np.array([Omega0, Gamma])

Omega0, Gamma, errOmega0, errGamma, fig4 = PHYS231.Phase(
    xData = ..., 
    yData = ..., 
    yErrors = ..., 
    start = start2,
    xlabel = "$\omega$", 
    ylabel = '$\phi$', 
    xUnits = 'krad/s', 
    yUnits = 'rad')
```

<!-- END QUESTION -->

In your lab notebook, you should compare this experimentally-measured values of $\omega_0$ and $\gamma$ to the values that you expected based on the values of the components used in your circuit.  Recall that $\omega_0 = 1/\sqrt{LC}$ and $\gamma=R/L$.  Keep in mind that your expected values of $\omega_0$ and $\gamma$ will also have uncertainties based on the uncertainties in the values of $R\pm\Delta R$, $L\pm\Delta L$, and $C\pm\Delta C$.

# <div style="background-color:rgba(255, 204, 255, 0.5); text-align:center; vertical-align: middle; padding:40px 0; margin-top:30px"><span style="color:rgba(102, 0, 204, 1);">Part 2 - Feedback and Submission</span></div>

<!-- BEGIN QUESTION -->

***
**<span style="color:blue">Question 2.1:</span>**  

We welcome your feedback on the PHYS 231 labs!  Please feel free to include any comments you have about this lab in the cell below.  Your comments will be taken into consideration when revising/improving the PHYS 231 labs.  You can suggest improvements, point out anything that was unclear, comment on the strengths and weaknesses of the lab, ...

This question is optional and will have no impact on your lab grade.

***
**<span style="color:blue">Answer 2.1:</span>**

[//]: # (Please do not delete this comment or anything above it.  Anything below this comment can be deleted.)  

Double click this cell to enter your comments.  This text can be deleted.

<!-- END QUESTION -->

---

To double-check your work, the cell below will rerun all of the autograder tests.

```python
grader.check_all()
```

***

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

Once you've completed this notebook:
- Save your work.
- Run 'grader.export()' to generate a .zip file containing all of the materials that you will submit.
- Download the generated .zip file.
- Upload the .zip file to the PHYS 121 Canvas gradebook.

Here is a <a href = "https://raw.githubusercontent.com/UBC-Okanagan-Physics-Labs/PHYS-121-images/main/general/gifs/Submission.gif">GIF</a> showing how these steps are completed.  Once your completed notebook has been uploaded to the Canvas gradebook, you're done!

```python
# Save your notebook first, then run this cell (place your cursor in the cell and then hit Shift + Enter) to export 
# your submission.  Uploaded the .zip file that is generated to the gradebook of the Canvas shell for the PHYS 121 lab.
grader.export()
```

# <div style="background-color:rgba(255, 204, 255, 0.5); text-align:center; vertical-align: middle; padding:40px 0; margin-top:30px"><span style="color:rgba(102, 0, 204, 1);">Part 3 - Playground (optional)</span></div>

Feel free to add as many cells as you like below and use them as a playground for further independent investigations.  These cells won't be graded, so feel free to use them in any way that you like.  

```python
# Here's an empty code cell that you can use.
```

```python
# Here's another empty code cell that you can use.
```

```python
# Here's yet another empty code cell that you can use.
# If you need more, you can add cells using the '+' icon in the menu bar at to the top of the screen.
```

### <div style="background-color:rgba(255, 255, 102, 0.5); text-align:left; padding:20px 0; margin-top:20px">$\quad$Appendix A &ndash; The Scatter Function...</div>

The function for generating scatter plots is called as follows:
```python
PHYS231.Scatter(xData, yData, yErrors = [], xlabel = 'x-axis', ylabel = 'y-axis', xUnits = '', yUnits = '', fill = False, show = True)
```
The 'xData' and 'yData' inputs are required, all other arguments are optional with default values set.  The function returns the a single output (the formatted plot):
```python
fig
```

The function will do a simple scatter plot if no 'yError' are included.  It will included error bars if 'yErrors' are passed to the function.  The 'fill' and 'show' arguments should generally be false.  

If the 'xData' list is empty, as in:
```python
xData = []
```
then the x-axis will be the trial number.  For example, is 'xData' is empty and 'yData' has 10 elements, the x-axis will span 1 to 10.

### Scatter Example Implmentation
The code block below shows an implementation of 'Scatter'.
```python
import PHYS231
theta = [10, 20, 30] # degrees
T = [2.02, 1.95, 2.13] # s
errT = [0.02, 0.02, 0.03] # s
fig = PHYS231.Scatter(theta, T, errT, 'initial angle' , 'period', 'degrees', 's')
```

If you're interested in generating your own plots with customized formatting, see the following Python-based plotting tutorial: https://cmps-people.ok.ubc.ca/jbobowsk/Python/html/Jupyter%20Basic%20Plots.html.
***

<img src="https://raw.githubusercontent.com/UBC-Okanagan-Physics-Labs/PHYS-121-images/main/general/images/ubc-logo-full.jpg" width="500"/>

Last update: October 12, 2023

Tests for name_and_student_number (points: None)
```
>>> isinstance(last_name_1, str)
True
```
failure_message: last_name_1 must be a string (i.e. surround your answer by either single or double quotes).
```
>>> isinstance(first_name_1, str)
True
```
failure_message: first_name_1 must be a string (i.e. surround your answer by either single or double quotes).
```
>>> isinstance(student_number_1, str)
True
```
failure_message: student_number_1 must be a string (i.e. surround your answer by either single or double quotes).
```
>>> test = True
>>> if isinstance(student_number_1, str):
...     test = len(student_number_1) == 8
>>> test
True
```
failure_message: student_number_1 must be a string of exactly 8 characters.
```
>>> test = True
>>> if isinstance(student_number_1, str):
...     test = student_number_1.isdigit()
>>> test
True
```
failure_message: student_number_1 must contain only interger digits (0, 1, 2, 3, ..., 9)
```
>>> isinstance(last_name_2, str) or last_name_2 == ...
True
```
failure_message: last_name_2 must be a string (i.e. surround your answer by either single or double quotes).
```
>>> isinstance(first_name_2, str) or first_name_2 == ...
True
```
failure_message: first_name_2 must be a string (i.e. surround your answer by either single or double quotes).
```
>>> isinstance(student_number_2, str) or student_number_2 == ...
True
```
failure_message: student_number_2 must be a string (i.e. surround your answer by either single or double quotes).
```
>>> test = True
>>> if isinstance(student_number_2, str):
...     test = len(student_number_2) == 8
>>> test
True
```
failure_message: student_number_2 must be a string of exactly 8 characters or be left as ...
```
>>> test = True
>>> if isinstance(student_number_2, str):
...     test = student_number_2.isdigit()
>>> test
True
```
failure_message: student_number_2 must contain only interger digits (0, 1, 2, 3, ..., 9) or be left as ...


Tests for q1.1 (points: None)
```
>>> len(frequency) == len(V0)
True
```
failure_message: Your frequency and V0 lists are different lengths.
```
>>> len(frequency) == len(errV0)
True
```
failure_message: Your frequency and errV0 lists are different lengths.
```
>>> len(frequency) == len(VR)
True
```
failure_message: Your frequency and VR lists are different lengths.
```
>>> len(frequency) == len(errVR)
True
```
failure_message: Your frequency and errVR lists are different lengths.
```
>>> len(frequency) == len(dt)
True
```
failure_message: Your frequency and dt lists are different lengths.
```
>>> len(frequency) == len(errdt)
True
```
failure_message: Your frequency and errdt lists are different lengths.
```
>>> len(V0) == len(errV0)
True
```
failure_message: Your V0 and errV0 lists are different lengths.
```
>>> len(V0) == len(VR)
True
```
failure_message: Your V0 and VR lists are different lengths.
```
>>> len(V0) == len(errVR)
True
```
failure_message: Your V0 and errVR lists are different lengths.
```
>>> len(V0) == len(dt)
True
```
failure_message: Your V0 and dt lists are different lengths.
```
>>> len(V0) == len(errdt)
True
```
failure_message: Your V0 and errdt lists are different lengths.
```
>>> len(errV0) == len(VR)
True
```
failure_message: Your errV0 and VR lists are different lengths.
```
>>> len(errV0) == len(errVR)
True
```
failure_message: Your errV0 and errVR lists are different lengths.
```
>>> len(errV0) == len(dt)
True
```
failure_message: Your errV0 and dt lists are different lengths.
```
>>> len(errV0) == len(errdt)
True
```
failure_message: Your errV0 and errdt lists are different lengths.
```
>>> len(VR) == len(errVR)
True
```
failure_message: Your VR and errVR lists are different lengths.
```
>>> len(VR) == len(dt)
True
```
failure_message: Your VR and dt lists are different lengths.
```
>>> len(VR) == len(errdt)
True
```
failure_message: Your VR and errdt lists are different lengths.
```
>>> len(errVR) == len(dt)
True
```
failure_message: Your errVR and dt lists are different lengths.
```
>>> len(errVR) == len(errdt)
True
```
failure_message: Your errVR and errdt lists are different lengths.
```
>>> len(dt) == len(errdt)
True
```
failure_message: Your dt and errdt lists are different lengths.
```
>>> N = 25
>>> len(frequency) >= N and len(V0) >= N and len(errV0) >= N and len(VR) >= N and len(errVR) >= N and len(dt) >= N and len(errdt) >= N
True
```
failure_message: A minimum of 35 measurements are required.


Tests for q1.2 (points: None)
```
>>> len(omega) != 0
True
```
failure_message: The omega array is empty.
```
>>> all(np.isclose(omega, 2*np.pi*frequency, rtol = 0.01))
True
```
failure_message: Double check your calculation of omega.  Remember that, in Python, the syntax for pi is np.pi.
```
>>> len(Ratio) != 0
True
```
failure_message: The Ratio array is empty.
```
>>> all(np.isclose(Ratio, VR/V0, rtol = 0.01))
True
```
failure_message: Double check your calculation of Ratio.
```
>>> len(errRatio) != 0
True
```
failure_message: The Ratio array is empty.
```
>>> all(np.isclose(errRatio, np.sqrt((errVR/V0)**2 + (VR*errV0/V0**2)**2), rtol = 0.01))
True
```
failure_message: Double check your calculation of Ratio.


Tests for q1.6 (points: None)
```
>>> len(Phi) != 0
True
```
```
>>> all(np.isclose(Phi, omega*1e3*dt*1e-6, rtol = 0.01))
True
```
failure_message: Double check your calculation of Phi.
```
>>> len(errPhi) != 0
True
```
```
>>> all(np.isclose(errPhi, omega*1e3*errdt*1e-6, rtol = 0.01))
True
```
failure_message: Double check your calculation of errPhi.
